Bifunctional Architecture Enables Substrate Catalysis and Channeling in Paracoccus TMAO Demethylase

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

Thach et al. report on the structure and function of trimethylamine N-oxide demethylase (TDM). They identify a novel complex assembly composed of multiple TDM monomers and obtain high-resolution structural information for the catalytic site, including an analysis of its metal composition, which leads them to propose a mechanism for the catalytic reaction.

In addition, the authors describe a novel substrate channel within the TDM complex that connects the N-terminal Zn2+-dependent TMAO demethylation domain with the C-terminal tetrahydrofolate (THF)-binding domain. This continuous intramolecular tunnel appears highly optimized for shuttling formaldehyde (HCHO), based on its negative electrostatic properties and restricted width. The authors propose that this channel facilitates the safe transfer of HCHO, enabling its efficient conversion to methylenetetrahydrofolate (MTHF) at the C-terminal domain as a microbial detoxification strategy. Experimental data that shows an involvement of TDM in the reaction of HCHO with THF is less convincing.

Strengths:

The authors provide convincing high-resolution cryo-EM structural evidence (up to 2 Å) revealing an intriguing complex composed of two full monomers and two half-domains. They further present evidence for the metal ion bound at the active site and articulate a hypothesis for the catalytic cycle. Substantial effort is devoted to optimizing and characterizing enzyme activity, including detailed kinetic analyses across a range of pH values, temperatures, and substrate concentrations. Furthermore, the authors validate their structural insights through functional analysis of active-site point mutants.

In addition, the authors identify a continuous channel for formaldehyde (HCHO) passage within the structure and support this interpretation through molecular dynamics simulations. These analyses suggest an exciting mechanism of specific, dynamic, and gated channelling of HCHO. This finding is particularly appealing, as it implies the existence of a unique, completely enclosed conduit that may be of broad interest, including potential applications in bioengineering.

Weaknesses:

Although the idea of an enclosed channel for HCHO is compelling, the experimental evidence supporting enzymatic assistance in the reaction of HCHO with THF is less convincing. The linear regression analysis shown in Figure 1C demonstrates a THF concentration-dependent decrease in HCHO; however, it is well established that HCHO and THF can react spontaneously in a non-enzymatic manner, raising the possibility that the observed effect does not require enzymatic involvement. I appreciate the authors' clarification that the data in Figure 1 were not intended to demonstrate enzymatic channelling or catalytic involvement in the HCHO-THF reaction, and that the assay does not distinguish between changes in HCHO production and downstream consumption. However, the statement "these findings show that TDM carries out two linked reactions: TMAO demethylation at one active site, and the HCHO produced can condense with THF at the C-terminal domain, connecting TMAO breakdown to one-carbon metabolism" (page 2) still implies a mechanistic and functional coupling that is not supported by the presented data and appears inconsistent with the authors' clarification. In light of this, I recommend revising this statement to avoid implying mechanistic or functional coupling between the two reactions unless additional experimental evidence is provided.

We thank the reviewer for this clarification. We have revised as per recommendation (page 2).

“Overall, these findings suggest that TDM-mediated TMAO demethylation generates HCHO, which can subsequently react with THF, potentially linking TMAO breakdown to one-carbon metabolism.”

Overall, the authors were successful in advancing our structural and functional understanding of the TDM complex. They suggest an interesting oligomeric complex composition which should be investigated with additional biophysical techniques.

Additionally, they provide an intriguing hypothesis for a new type of substrate channelling. Additional kinetic experiments focusing on HCHO and THF turnover by enzymatic proximity effects would strengthen this potentially fundamental finding. If this channelling mechanism can be supported by stronger experimental evidence, it would substantially advance our understanding and knowledge of biologic conduits and enable future efforts in the design of artificial cascade catalysis systems with high conversion rate and efficiency, as well as detoxification pathways.

Reviewer #2 (Public review):

Summary:

The manuscript reports a cryo-EM structure of TMAO demethylase from Paracoccus sp. This is an important enzyme in the metabolism of trimethylamine oxide (TMAO) and trimethylamine (TMA) in human gut microbiota, so new information about this enzyme would certainly be of interest.

Strengths:

The cryo-EM structure for this enzyme is new and provides new insights into the function of the different protein domains, and a channel for formaldehyde between the two domains.

Weaknesses:

(1) The proposed catalytic mechanism in this manuscript does not make sense. Previous mechanistic studies on the Methylocella silvestris TMAO demethylase (FEBS Journal 2016, 283, 3979-3993, reference 7) reported that, as well as a Zn2+ cofactor, there was a dependence upon non-heme Fe2+, and proposed a catalytic mechanism involving deoxygenation to form TMA and an iron(IV)-oxo species, followed by oxidative demethylation to form DMA and formaldehyde.

In this work, the authors do not mention the previously proposed mechanism, but instead just say that elemental analysis "excluded iron". This is alarming, since the previous work has a key role for non-heme iron in the mechanism. The elemental analysis here gives a Zn content of about 0.5 mol/mol protein (and no Fe), whereas the Methylocella TMAO demethylase was reported to contain 0.97 mol Zn/mol protein, and 0.35-0.38 mol Fe/mol protein. It does, therefore, appear that their enzyme is depleted in Zn, and the absence of Fe impacts on the mechanism, as explained below.

The proposed catalytic mechanism in this manuscript, I am sorry to say, does not make sense, for several reasons:

(i) Demethylation to form formaldehyde is not a hydrolytic process; it is an oxidative process (normally accomplished by either cytochrome P450 or non-heme iron-dependent oxygenase). The authors propose that a zinc (II) hydroxide attacks the methyl group, which (a) is unprecedented, (b) even if it were possible, would generate methanol, not formaldehyde.

(ii) The amine oxide is proposed to deoxygenate, with hydroxide appearing on the Zn - unfortunately, amine oxide deoxygenation is a reductive process, for which a reducing agent is needed, and Zn2+ is not a redox-active metal ion;

(iii) The authors say "forming a tetrahedral intermediate, as described for metalloprotease," but zinc metalloproteases attack an amide carbonyl to form an oxyanion intermediate, whereas in this mechanism, there is no carbonyl to attack, so this statement is just wrong.

So on several counts the proposed mechanism cannot be correct. Some redox cofactor is needed in order to carry out amine oxide deoxygenation, and Zn2+ cannot fulfil that role. Fe2+ could do, which is why the previously proposed mechanism involving an iron(IV)-oxo intermediate is feasible. But the authors claim that their enzyme has no Fe. If so then there must be some other redox cofactor present. Therefore, the authors need to re-analyse their enzyme carefully and look either for Fe or for some other redox-active metal ion, and then provide convincing experimental evidence for a feasible catalytic mechanism. As it stands the proposed catalytic mechanism is unacceptable.

Revised version. The authors have essentially not changed the proposed mechanism. They have removed the reference to zinc metalloproteases, but still propose a mechanism mediated only by Zn2+. As explained above, attack by zinc (II) hydroxide is unprecedented and would generate methanol, not formaldehyde, and amine deoxygenation is a reductive process that cannot be fulfilled by Zn2+. So the proposed mechanism is still not feasible at all. The authors now say that "oxidative chemistry....remains unresolved", I'm sorry, but that is not acceptable.

I have urged the authors to re-examine the metal content of their enzyme, In the Supporting Information (Figure S5) they give ICPMS data that indicates a Zn stoichiometry of 0.5 mol Zn/mol protein, and Fe is not detected. Have the authors analysed for other redox active metals? The authors say that there is no evidence for any other metal binding site, but there is only 50% occupancy of Zn in their protein, so could there be a different metal ion present in place of Zn in the other 50% of the protein, that accounts for the observed activity?

Since there is clearly a major discrepancy here, the onus is on the authors to explain the discrepancy, rather than just returning with the same data. For example, they could treat the enzyme with EDTA to remove all metals (and check the treated enzyme by ICPMS), and then add different metal ions to test activity with different metals (could even titrate with different molar equivalents of metal ions). They could then test a range of different redox-active metal ions.

We have re-examined our data and repeated experiments on the reviewer's opinion. We have repeated the IC-PMS several times with different preps, including full scans (data presented). Our enzyme is active, but no iron signal is detected. Moreover, the experimentally determined structure does not support the presence of a non-heme iron-binding site (Bugg TDH, Ramaswamy S., doi:10.1016/j.cbpa.2007.12.007, and other papers). More detailed response in the recommendations to authors.

(2) Given the metal content reported here, it is important to be able to compare the specific activity of the enzyme reported here with earlier preparations. The authors have now done this in the revised version.

(3) The consumption of formaldehyde to form methylene-THF is potentially interesting, but the authors say "HCHO levels decreased in the presence of THF", which could potentially be due to enzyme inhibition by THF. Is there evidence that this is a time-dependent and protein-dependent reaction? Not yet addressed.

We thank the reviewer for this important point. At present, we have not performed detailed time-dependent or protein-dependent analyses to determine whether the observed decrease in HCHO levels in the presence of THF reflects enhanced downstream consumption or indirect effects, such as inhibition of TDM activity by THF. We acknowledge that further kinetic and protein-dependence studies will be important directions for future work.

Also in Figure 1C, HCHO reduction (%) is not very helpful, because we don't know what concentration of formaldehyde is formed under these conditions; it would be better to quote in units of concentration, rather than %. This point has been addressed by the authors in the revised version.

(4) Has this particular TMAO demethylase been reported before? It's not clear which Paracoccus strain the enzyme is from; the Experimental Section just says "Paracoccus sp.", which is not very precise. There has been published work on the Paracoccus PS1 enzyme, is that the strain used? Details about the strain are needed, and the accession for the protein sequence. Addressed in the revised version.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

As noted above, there is still a major problem with the proposed mechanism not being feasible, and remaining questions about the presence or absence of a redox-active metal ion in their enzyme. They should:

(1) Re-examine for other metal ions (apart from Zn and Fe) using ICPMS. The redox metal ion could, in theory, be some other transition metal ion.

(2) Seek evidence for the role of metal ions in the activity of this enzyme, for example, by treating with EDTA to remove metal ions, and then adding different metal ions, to correlate activity with a particular metal ion.

We thank the reviewer for these valuable suggestions regarding the metal identity and its functional role in TDM activity. We also acknowledge the reviewer’s concerns regarding the catalytic mechanism. In response, we have substantially revised Scheme 1 and the associated Discussion text to focus on the observed interactions of the substrate (TMAO) and products (DMA and HCHO) within the Zn2+-containing active site, rather than proposing a detailed catalytic mechanism that is not fully supported by the current data.

(1) We repeated the ICP–MS analysis using a wide-range full-scan survey to examine the presence of additional metal-associated isotopes beyond Zn and Fe. The corresponding experimental details have been added to the revised ICP–MS Methods section (Page 9). Full-scan ICP–MS profiling of purified TDM detected Zn as the predominant associated metal species (Author response image 1A). In contrast, signals corresponding to Fe and other transition metals were either undetectable or present only at trace levels comparable to, or lower than, those observed in the digested HNO3 solution control. To further validate this observation, we performed targeted ICP–MS quantification for both Zn and Fe on the same purified samples. These measurements confirmed that Fe was below the detection threshold, whereas Zn was consistently detected at an approximate ratio of 0.5 Zn2+ per protein monomer (Author response image 1B, C, Figure S5).

The observed 0.5 Zn2+-to-protein stoichiometry is consistent with the previously discussed 2 full-length + 2 half-domain (2+2½) assembly. In this complex, only the intact core domains retain the complete metal-binding motif, whereas the truncated half-domains lack the Zn2+-binding region. Consequently, only two metal-binding sites are expected per assembled complex, in agreement with the ICP–MS measurements. We additionally note that Zn2+ was not intentionally supplemented during purification. Based on the current cryo-EM and biochemical data, both metal-binding sites in the full-length subunits appear similarly occupied, with no evidence for asymmetric metal loading.

Importantly, the previously published FEBS Journal model proposed an Fe2+-binding site based on metal analysis and homology modeling rather than direct experimental determination. The residues implicated in Fe2+ binding are well resolved in our experimental maps and do not define a metal-coordination environment compatible with a second mononuclear metal-binding site. Consistent with the ICP–MS results, the experimental structure provides no evidence of a second metal-binding site. Nevertheless, the enzyme remains catalytically active under these conditions.

(2) We attempted metal depletion experiments using EDTA treatment to evaluate the functional role of the bound metal ion. However, removal of metal ions resulted in rapid protein aggregation, preventing subsequent activity measurements. These observations suggest that the bound Zn2+ ion plays an important role in maintaining the structural integrity and stability of the TDM complex. While metal reconstitution experiments would be informative, the aggregation observed following metal depletion precluded a meaningful assessment of alternative metal ions in the current study.

Changes from the earlier version to the revised version

Abstract

Substrate channeling enhances efficiency and prevents toxicity by directing unstable intermediates between active sites. Trimethylamine N-oxide demethylase (TDM) degrades trimethylamine N-oxide (TMAO) to dimethylamine and formaldehyde (HCHO), but the fate of HCHO has remained unclear. We report cryo-EM structures of TDM in apo, substrate-, and product-bound states that reveal a previously unknown channeling pathway. Combined structural, biochemical, and target molecular dynamics analyses show that HCHO is generated in a catalytic core and guided through a tunnel to a remote tetrahydrofolate (THF)-binding site, where it forms methylene-THF. Thus, TDM emerges as a bifunctional enzyme that unites TMAO demethylation with one-carbon transfer, providing a mechanistic explanation for its role in metabolic efficiency and detoxification.

Introduction

Substrate channeling occurs in enzymatic reactions, where intermediates are transferred directly from one enzyme’s active site to another, bypassing their release into the cellular environment.[1,2] This mechanism improves catalytic efficiency by minimizing intermediate loss, while simultaneously protecting cells from the damaging effects of unstable or toxic species.[1,2,3] Classic examples include glycolytic and amino acid metabolic enzymes, where covalent or noncovalent tunnels ensure safe and rapid passage of intermediates. [1,2,3] Despite its importance, structural evidence for substrate channeling remains limited, particularly in enzymes that couple detoxification with energy metabolism.

Trimethylamine (TMA) and its oxidized derivative trimethylamine N-oxide (TMAO) are abundant small molecules present in both marine organisms and the human gut microbiome.[4,5] TMA arises from microbial metabolism of dietary choline, carnitine, and other animal-derived nutrients, and is subsequently oxidized to TMAO by host flavin-containing monooxygenases.[4],[5] These metabolites are physiologically and clinically significant; elevated TMAO levels have been associated with cardiovascular disease and gastrointestinal cancer.[4,5] Conversely, in marine bacteria and certain gut microbes, TMAO serves as both an energy source and a carbon donor, fueling specialized catabolic pathways that convert it into downstream metabolites.[4,5,6]

A key enzyme in this process is trimethylamine N-oxide demethylase (TDM), which catalyzes the cleavage of a methyl group from TMAO to generate dimethylamine (DMA) and formaldehyde (HCHO).[6,7] Early biochemical studies revealed that TDM is a two-domain protein: an N-terminal domain that mediates TMAO demethylation, and a C-terminal domain structurally related to tetrahydrofolate (THF)-dependent enzymes, including the glycine cleavage T protein and dimethylsulfoniopropionate demethylase.[7,8,9] This homology led to the hypothesis that the C-terminal domain could utilize the HCHO intermediate to form methylenetetrahydrofolate (MTHF), thereby coupling detoxification with one-carbon transfer metabolism.[7,9] However, the molecular details of this bifunctional activity, how reactive formaldehyde is channeled and presented to the THF-binding site, have remained unresolved.

Here, we report cryo-EM structures of Paracoccus TDM captured in apo, substrate-bound, and product-bound states, providing the first high-resolution view of its dual-domain architecture (local resolution of 2.0 Å). Structural analysis revealed that the TMAO- and THF-binding pockets are separated by approximately 60 Å, raising the question of how formaldehyde is safely transferred. Tunnel mapping combined with target molecular dynamics simulations identified a continuous, negatively charged channel linking the two active sites, suggesting a structural basis for substrate channeling. Complementary biochemical assays and mass spectrometry confirmed that HCHO generated at the Core catalytic site is efficiently converted to MTHF at the C-terminal domain. Together, these results establish TDM as a bifunctional enzyme that unites TMAO demethylation with formaldehyde detoxification and one-carbon transfer, highlighting a microbial strategy to enhance metabolic efficiency while mitigating toxicity.

Results and discussion

Optimization of TDM Activity and Evidence for Bifunctionality

Earlier version: We first optimized purification and assay conditions for TDM by systematically varying salt concentration, pH buffer, and temperature. These analyses revealed a broad range, with the highest activity between pH 6.0 and 7.0 at 150 mM NaCl and 35-45 °C (Figure S1A–D). Under these conditions, TDM robustly demethylated TMAO, yielding HCHO with a Vmax of 156 nmol/min/mg and a Km of 1.33 mM (Figure 1A). Notably, these kinetic parameters indicate higher catalytic activity than the previously reported TDM from Methylocella silvestris, which exhibits a Km of 3.88 mM and a Vmax of 14.61 nmol/min/mg.[7] HPLC-MS confirmed the production of DMA as a co-product (Figure 1B), consistent with prior reports.[7] Notably, HCHO concentrations decreased upon addition of THF (Figure 1C), suggesting that HCHO may be further processed into MTHF (Figure 1D). Overall, these findings show that TDM carries out two linked reactions: TMAO demethylation at one active site, and the HCHO produced can condense with THF at the C-terminal domain, connecting TMAO breakdown to one-carbon metabolism.

Revised version: We first optimized purification and assay conditions for TDM by systematically varying salt concentration, pH buffer, and temperature. These analyses revealed a broad range, with the highest activity between pH 6.0 and 7.0 at 150 mM NaCl and 35-45 °C (Figure S1A–D). Under these conditions, TDM robustly demethylated TMAO, yielding HCHO with a Vmax of 156 nmol/min/mg and a Km of 1.33 mM (Figure 1A). Notably, these kinetic parameters indicate higher catalytic activity than the previously reported TDM from Methylocella silvestris, which exhibits a Kmof 3.88 mM and a Vmaxof 14.61 nmol/min/mg.[7] HPLC-MS confirmed the production of DMA as a co-product (Figure 1B), consistent with prior reports.[7] Notably, HCHO concentrations decreased upon addition of THF (Figure 1C), suggesting that HCHO may be further processed into MTHF (Figure 1D). Overall, these findings suggest that TDM-mediated TMAO demethylation generates HCHO, which can subsequently react with THF, potentially linking TMAO breakdown to one-carbon metabolism.

TDM catalyzes two sequential reactions.

(A) Steady-state kinetic analysis of TDM - producing HCHO, in response to varying concentrations of the substrate TMAO (n=3). (B) Representative HPLC-MS spectrum of DMA product from TDM:TMAO reaction. The standard curve for DMA is shown in green, and that for TMAO is shown in orange. The red trace shows the product of the reaction when TMAO is added to TDM, indicating the formation of DMA. Counts in the Y-axis. (C) Linear regression analysis demonstrates a concentration-dependent decrease in HCHO production when THF is included in the reaction mixture, suggesting THF-dependent consumption or conversion of HCHO. (D) Representative HPLC-MS confirming the formation of MTHF adduct as a product of the TDM:TMAO:THF reaction. The red, gray, and green lines represent the standard curves for THF, MTHF, and the reaction product when TMAO and THF were present with TDM, respectively.

Cryo-EM Structures Reveal an Unusual Oligomeric Assembly

To elucidate the catalytic mechanism, we determined cryo-EM structures of apo TDM, a product-bound state (TMAO-soaked wild type), and a substrate-trapped mutant (D220A/D367A in complex with TMAO) (Figure 2A–B, S2, Table 1). TDM assembles into an unexpected 2 + 2½ oligomeric architecture (Figure 2A–C). Two full-length subunits form a highly symmetric core dimer (RMSD = 0.12 Å over 787 Cα atoms), while two peripheral half-subunits consist solely of C-terminal THF-binding domains (Figure 2C). The core dimer interface is stabilized by extensive hydrogen-bonding and hydrophobic interactions (Figure S3).

Overall architecture of TDM determined by cryo-electron microscopy.

(A) Representative cryo-electron microscopy (EM) maps. (B) Overall structure of the complex, with each subunit colored differently. TDM adopts a 2 + 21/2 complexation. The two half-domains of the complex are the C-terminal domains. (C) (C) Schematic illustration of TDM complexation. TDM monomers with N-terminal, Core, and C-terminal domains are colored forest, orange, and blue, respectively. The same coloring is used in panels A and B. (D) Close-up view of the 3Cys:Zn2+ binding motif. (E) Close-up view of the interactions between DMA and HCHO products with the 3Cys:Zn2+ motif. (F) The TMAO substrate is trapped in the 3Cys:Zn2+ active site of the double D220A/D367A mutant. The density maps are shown and contoured at 3.0 σ with a carve of 2 Å. Bound small molecules are shown as ball-and-stick representations, amino acid residues as sticks, and hydrogen bonds as dashed lines. (G) Relative activity of TDM variants.

Each catalytic core domain contains a conserved 3Cys:Zn2+ motif (C285, C301, C365) essential for protein stability (Figure 2D, S4). Substitution of any of these cysteines with alanine led to aggregation, underscoring their structural importance. Elemental analysis (EDS and ICP–MS) confirmed zinc incorporation but excluded iron, yielding a Zn: protein ratio of ∼1:2 (Figure S5A, B). This motif is analogous to Zn2+ centers in cytidine deaminases,[10,11] alcohol dehydrogenases, and other Zn metalloenzymes.[12,13]

Zn2+-Dependent Organization at the Core Catalytic Domain

Earlier version: In TDM, Zn2+ occupies a single, well-defined site within the core catalytic domain and plays a central non-redox role in substrate binding and organization of the active-site environment. Cryo-EM analysis of the TDM: TMAO complex reveals clear densities corresponding to the reaction products dimethylamine (DMA) and formaldehyde (HCHO) within the Zn2+-containing active site (Figure 2E), supporting its direct involvement in substrate turnover. Zn2+ coordinates the TMAO substrate via its oxygen atom and contributes to proper substrate positioning in concert with conserved residues D220, F327, and D367 (Figure 2E).

Revised version: In TDM, Zn2+ occupies a single, well-defined site within the core catalytic domain and plays a central non-redox role in substrate binding and organization of the active-site environment. Cryo-EM analysis of the substrate-bound TDM complex reveals TMAO positioned within the Zn2+-containing active site through coordination of its oxygen atom with the metal center, while the surrounding conserved residues D220, F327, and D367 contribute to substrate recognition and stabilization (Figure 2D-F). These interactions define the substrate-bound state of the catalytic pocket and are summarized in Scheme 1, which illustrates the experimentally observed interactions between TMAO and the Zn2+-containing active site.

Earlier version: The cryo-EM structure of the TDM:DMA: HCHO complex shows both products retained within the active site, whereas in the D220A/D367A: TMAO mutant structure, the substrate remains trapped without substantial conformational changes relative to the wild-type enzyme (Figure 2E, F). Consistent with these observations, the F327A and D220A variants exhibit markedly reduced enzymatic activity, underscoring their importance in substrate recognition and active-site organization (Figure 2G). In contrast, mutation of D367, which participates in hydrogen bonding with the DMA product, results in a more modest (∼10%) reduction in activity, consistent with a role in product stabilization rather than substrate binding or activation (Figure 2E–G).

Revised version: The cryo-EM structure of the TDM:DMA:HCHO complex shows both products retained within the active site, whereas in the D220A/D367A:TMAO mutant structure, the substrate remains trapped without substantial conformational changes relative to the wild-type enzyme (Figure 2E, F). Consistent with these observations, the F327A and D220A variants exhibit markedly reduced enzymatic activity, underscoring their importance in substrate recognition and active-site organization (Figure 2G). In contrast, mutation of D367, which participates in hydrogen bonding with the DMA product, results in a more modest (∼10%) reduction in activity, consistent with a role in product stabilization rather than substrate binding or activation (Figure 2E–G). Together with the substrate-bound structure, these product-associated interactions complete the structural framework summarized in Scheme 1, illustrating how the Zn2+-containing active site accommodates both the substrate and the reaction products during TMAO turnover. While the detailed chemical mechanism remains to be established, these structures provide a framework for understanding substrate recognition, product stabilization, and active-site organization in TDM.

Added in the revised version: Interactions of the substrate (TMAO) and reaction products within the Zn2+-containing active site of TDM.

Added in the revised version: TMAO is coordinated within the active site through interactions with Zn2+ and surrounding conserved residues. The proposed product-binding states illustrate interactions of dimethylamine (DMA) and formaldehyde (HCHO) with residues lining the active-site pocket following turnover.

Removed in the revised version: Based on our structural and biochemical data, we propose a structurally informed working model for TMAO turnover by TDM (Scheme 1). In this model, Zn2+ plays a non-redox role by polarizing the O–H bond of the bound hydroxyl, thereby lowering its pKa. The D220 carboxylate functions as a general base, abstracting the proton to generate a hydroxide nucleophile. This hydroxide then attacks the electrophilic N-methyl carbon of TMAO, forming a tetrahedral carbinolamine (hemiaminal) intermediate. Subsequent heterolytic cleavage of the C–N bond leads to the release of HCHO. D220 then switches roles to act as a general acid, donating a proton to the departing nitrogen, thereby facilitating product release and regenerating the active site. This sequence allows a new water molecule to rebind Zn2+, enabling subsequent catalytic turnovers. This proposed pathway is consistent with prior mechanistic studies, in which water addition to the azomethine carbon of a cationic Schiff base generates a carbinolamine intermediate, followed by a rate-limiting breakdown to yield an amino alcohol and a carbonyl compound, in the referenced case, an aldehyde.[14]

Earlier version: In particular, we observe no structural or biochemical support for an additional iron-binding site. Previous studies of TDM from Methylocella silvestris proposed an iron-dependent mechanism based on metal analysis and mutagenesis, in the absence of structural information. Distinct binding sites for Zn2+ and non-heme Fe2+ were not experimentally determined.[7,15] In contrast, our structural analyses consistently reveal a single metal-binding site occupied by Zn2+. Even when TDM was expressed in LB medium supplemented with 0.5 mM Fe(NH4)2SO4 (Figures S6A and S6B), we did not observe a second iron-binding site. Thus, Zn2+ appears to function primarily as an organizational cofactor at the core catalytic domain, while the source of the oxidative chemistry required for demethylation remains unresolved.

Revised version: Notably, we found no structural or biochemical evidence for an additional iron-binding site in TDM. Previous studies of TDM from Methylocella silvestris proposed the involvement of both Zn2+ and non-heme Fe2+ based on metal-content analyses and mutagenesis experiments; however, in the absence of experimentally determined structural information, the locations and coordination environments of these metals could not be directly resolved.[7,14] In contrast, our ICP-MS and cryo-EM analyses consistently identified a single Zn2+-binding site within the core catalytic domain. Furthermore, expression of TDM in LB medium supplemented with 0.5 mM Fe(NH4)2SO4 (Figures S6A and S6B) did not reveal detectable Fe incorporation or the presence of an additional metal-binding site. Collectively, these data support a single Zn2+-containing active site in our TDM preparation, and suggest that its metal composition differs from that proposed for the M. silvestris enzyme. While Zn2+ is structurally associated with substrate binding and organization of the active-site environment, the detailed chemical basis of TMAO demethylation remains to be established.

Removed in the revised version: Structurally informed working model for TDM-mediated TMAO turnover.

Removed in the revised version: TMAO binds at a single Zn2+-containing active site, where Zn2+ plays a non-redox role in substrate coordination and organization of the catalytic environment. Conserved residues contribute to substrate positioning and product stabilization. D220 acts as the catalytic base.

Architecture and Binding Properties at THF-Binding Domain

Earlier version: The C-terminal THF-binding domain is positioned on the intramolecular face and forms extensive electrostatic interactions that are essential for stabilizing the central dimer interface (Figure S7A, B). Structural comparison revealed high similarity to a THF-binding protein, with superposition onto a T protein (PDB ID: 1WOO) yielding an RMSD of 1.18 Å over 371 Cα atoms (Figure S7C).

Revised version: The C-terminal THF-binding domain is positioned on the intramolecular face of the TDM complex, and forms extensive electrostatic interactions that are essential for stabilizing the central dimer interface (Figure S7A, B).[15,16] Structural comparison revealed that this domain adopts a conserved folate-binding fold closely resembling those of previously characterized THF-dependent proteins. Superposition with the T protein of the glycine cleavage system (PDB ID: 1WOO) yielded an RMSD of 1.18 Å over 371 Cα atoms (Figure S7C), indicating a high degree of structural conservation despite limited sequence identity.[15] In addition, the overall architecture is consistent with the folate-binding domain of formiminotransferase-cyclodeaminase, another THF-dependent enzyme that utilizes a dedicated folate-binding module and has been proposed to facilitate intermediate transfer between catalytic domains.[16]

Cryo-EM density at the folate-binding site was consistent with that of THF binding, although resolution limitations and partial occupancy prevented complete ligand assignment (Figure 3A). Based on the structural homology, we modeled THF in the central cavity of the ring structure (Figure 3B). The pterin moiety of THF projected into the interior cavity, forming hydrophobic contacts with the aromatic folate-binding pocket (Figure 3B). Two positively charged residues, R480 and K782, located near the entrance of the binding funnel, likely engaged the polyglutamate tail via hydrogen bonding, complemented by π-stacking from F582 and F672 (Figure 3B). Similar binding modes have been observed in enzymes, such as T-proteins, aminotransferase domains, and dimethylglycine dehydrogenase.[17,18,19] ITC binding isotherms show a clear interaction between THF and TDM, with negligible heat change in the buffer control, confirming binding specificity. The dissociation constant (Kd) was determined to be approximately 149 nM (Figure 3C, S8).

THF binding to TDM.

(A) The THF-binding domain forms a cloverleaf of three subdomains (I-III), creating a central cavity that accommodates THF. EM density for THF is shown as a surface. (B) Close-up of the binding pocket highlighting hydrophobic and hydrogen-bonding residues that stabilize THF, colored according to their subdomain location in panel A.(C) ITC analysis confirms specific THF binding to TDM (blue), with negligible signal in buffer control (gray).

Substrate Channel for Formaldehyde Transfer

High-resolution cavity mapping of the TDM:DMA:HCHO: THF complex revealed a continuous intramolecular tunnel that connects the Zn2+-containing demethylase pocket to the distal THF-binding site (Figure 4A). This conduit begins at the formaldehyde release position within the catalytic pocket at the core domain (HA, HB) and terminates near the folate pterin ring at the C-terminal domain (TA, TB). The channel is predominantly lined with acidic residues and polar side chains, generating a strongly negative electrostatic potential (Figure 4A). Such electrostatics likely stabilize the partial negative oxygen of HCHO, guiding it through the tunnel while excluding bulk anions and minimizing leakage.

HCHO channels from the active site to the THF-binding site in TDM complex.

(A, B) Surface representations of the channels within the TDM complex colored according to either electrostatic potential or radius. (C) Tunnel analysis of a CG MD simulation suggests that HCHO in each chain can access the THF binding sites in chains A and B. The five top-ranked paths for chain A are shown in different colors. THF, DMA, and HCHO are depicted as spheres colored by the atom type. TA, TB, TC, and TD: THF molecules in chains A, B, C, and D, respectively; HA: chain A, HCHO.

Geometrically, the tunnel exhibits a radius ranging between 2.6 and 5.0 Å, with the narrowest bottleneck of 2.6 Å positioned just before the folate-binding cavity (Figure 4B). This matches closely with the van der Waals dimensions of formaldehyde (∼2.5 Å), suggesting evolutionary fine-tuning for selective transfer of this reactive one-carbon intermediate. Importantly, no alternative solvent-accessible exits were detected, indicating that HCHO is obligatorily channeled to the THF site rather than diffusing freely.

Coarse-grained target molecular dynamics simulations further supported this interpretation. Independent trajectories consistently identified two persistent tunnels, each connecting the Zn2+ demethylation sites (HA, HB) to the corresponding folate-binding sites (TA, TB) in the two protomers (Figure 4C, S9). Within these trajectories, formaldehyde molecules entered the channel from the catalytic site, transiently interacted with acidic residues and ordered water molecules, and passed through narrow gating residues that rearranged to allow transit (Figure 4C). This dynamic yet persistent tunnel suggests a gated channeling mechanism that balances selectivity with throughput.

Earlier version: The channeling mechanism in TDM bears resemblance to classical substrate-channeling systems but with distinct architectural solutions. In tryptophan synthase, an ∼25 Å hydrophobic tunnel delivers indole between two active sites,[20,21] whereas the glycine cleavage system employs a swinging lipoate arm on the H-protein to transfer intermediates between catalytic partners. [22,23] Dimethylglycine dehydrogenase also channels formaldehyde to THF, but does so through surface-proximal transfer rather than via a buried conduit,[9,19] while the PaaZ enzyme in phenylacetate catabolism uses electrostatic pivoting, in which a domain-swapped CoA arm swings the intermediate between domains, guided by conserved positive residues.[23] By contrast, TDM relies on a fully enclosed, negatively charged tunnel, a unique design optimized for formaldehyde, a small, highly diffusible, and cytotoxic intermediate.

Revised version: The channeling mechanism in TDM bears resemblance to classical substrate-channeling systems but with distinct architectural solutions. In tryptophan synthase, an ∼25 Å hydrophobic tunnel delivers indole between two active sites,[20,21] whereas the glycine cleavage system employs a swinging lipoate arm on the H-protein to transfer intermediates between catalytic partners.[22,23] Dimethylglycine dehydrogenase also channels formaldehyde to THF, but does so through surface-proximal transfer rather than via a buried conduit,[9,19] while the PaaZ enzyme in phenylacetate catabolism uses electrostatic pivoting, in which a domain-swapped CoA arm swings the intermediate between domains, guided by conserved positive residues.[22] By contrast, TDM relies on a fully enclosed, negatively charged tunnel, a unique design optimized for formaldehyde, a small, highly diffusible, and cytotoxic intermediate.

Role of the N-terminal domain in TDM assembly and function

The function of the N-terminal domain in TDM has remained unclear, and sequence alignment across homologs reveals low evolutionary conservation in this region (Figure 5A). Cryo-EM analysis revealed two distinct interaction modes between the N-terminal and C-terminal domains, including an unexpected interdomain interface between the N-terminal region and a peripheral THF-binding domain (Figure 5B, S10, Video S1). This interface appears to tether the flexible N-terminus to the THF-binding domain, forming a structural bridge that may regulate the conformational transitions required for complex assembly and efficient substrate transfer. The interaction is stabilized by aromatic stacking between P14 and W476, charge interaction, along with extensive hydrophobic contacts that promote proper domain packing and interdomain communication (Figure 5B, S9).

The N-terminal domain is involved in the structural stability and assembly of proteins.

(A) The evolutionary conservation of residue positions on the TDM structure is estimated based on phylogenetic connections among homologous sequences. (B) N-terminal domain was depicted as forest. Residues ranging from 1-20, 150-158 were highlighted. (C) Relative TDM activity of TDM mutants compared to that of WT variant.

Functional analyses further support this structural coupling. Deletion of the entire N-terminal region (Δ1–150) or the interfacial segment (Δ151–158) led to severe aggregation and loss of soluble protein, while removal of only the first 20 residues (Δ1–20) caused a modest (∼10%) reduction in catalytic turnover (Figure 5C). These results indicate that the N-terminal domain not only stabilizes the quaternary architecture but may also facilitate allosteric coupling between the catalytic and folate-binding domains. Such dynamic coupling likely ensures efficient handoff of formaldehyde through the internal tunnel, synchronizing demethylation and one-carbon transfer reactions within the bifunctional TDM complex.

Conclusion

Our results establish TDM as both a bifunctional enzyme and an integrated substrate-channeling system. By coupling TMAO demethylation with THF-dependent one-carbon transfer, TDM not only produces DMA and methylene-THF but also ensures immediate capture of formaldehyde, a highly reactive and cytotoxic intermediate. Structural, biochemical, and computational analyses reveal that this is achieved through a fully enclosed, negatively charged conduit that links the Zn2+ catalytic core to the folate-binding site, thereby enhancing catalytic efficiency and preventing metabolite leakage.

Beyond defining the molecular basis of TDM activity, these findings highlight a general design principle: substrate channeling can evolve distinct architectures, hydrophobic tunnels, swinging arms, or electrostatic conduits, tailored to the physicochemical properties of the intermediate. The TDM structure expands this paradigm by revealing a tunnel optimized for formaldehyde transfer, underscoring microbial strategies for detoxification and metabolic efficiency. More broadly, this work provides a framework for re-engineering channeling systems in biocatalysis, synthetic biology, and microbiome modulation, where controlling the fate of reactive intermediates is central to both metabolic innovation and therapeutic applications.

Experimental details

TDM Construction, Expression, and Purification

The gene encoding full-length TDM (1-787) (WP_26356686) from Paracoccus sp. DMF was optimized for E. coli codons and cloned into the pET21a(+) plasmid with either a C-terminal 6xHis tag or an N-terminal strepII tag using NdeI and XhoI restriction enzymes. Site-directed mutagenesis involving substitutions or deletions was performed using specific PCR primers (Universal, USA). Plasmid isolation and purification from E. coli were performed using a DNA spin kit (Qiagen, Germany), and the identity of the DNA constructs was confirmed by sequencing (Genewiz, USA).

The TDM protein variants were expressed in E. coli BL21 (DE3) cells cultured in LB medium at 37 °C. Induction was achieved with 0.25 mM isopropyl β-D-1-thiogalactopyranoside (IPTG) at 18 °C for 16 h. Protein purification was performed by affinity chromatography using either Ni-NTA resin (Qiagen, Germany) or a strep-tag column (Cytiva, USA) in buffer A (20 mM Tris-HCl, pH 7.5, 250 mM NaCl, 1 mM TCEP, 300 mM imidazole, and 5% glycerol). Further purification steps included passage through a 2 × 5 mL HiTrap Q column (GE Healthcare, USA) using buffer B (20 mM Tris-HCl pH 7.5, 200 mM NaCl, 1 mM TCEP, and 5% glycerol) for equilibration, and buffer C (20 mM Tris-HCl pH 7.5, 1000 mM NaCl, 1 mM TCEP, and 5% glycerol) for elution. Final purification was achieved using a Superdex 200 increase 10/300 GL column (Cytiva, USA) in buffer D (20 mM Tris-HCl, pH 7.5, 150 mM NaCl, and 5% glycerol). The purity of the relevant fractions was assessed by SDS-PAGE.

TDM Activity Assay

Earlier version: Enzyme activity was assessed by measuring HCHO production from TMAO breakdown as described previously with minor modifications.[24] Briefly, enzymatic activity was evaluated in multiple buffering systems spanning a range of pH values. Maximal activity was observed in 10 mM MES buffer (pH 6.0) supplemented with 150 mM NaCl; this buffer was subsequently used for all kinetic analyses. Standard activity assays were performed in triplicate at room temperature. Each 50 μL reaction contained 2.5 μg of purified TDM in MES buffer and was initiated by the addition of TMAO to a final concentration of 10 mM. Reactions were incubated for 10 min, a time point confirmed to be within the linear range based on preliminary time-course experiments. For HCHO detection, 10 μL of the reaction mixture was combined with 25 μL of freshly prepared 0.2% (w/v) Purpald reagent dissolved in 1 M NaOH and 215 μL of Milli-Q water in a 96-well plate. After incubation for 20 min at room temperature, absorbance at 540 nm was measured using a SpectraMax G5 microplate reader. Formaldehyde concentrations were calculated from a standard curve generated using analytical-grade formaldehyde (0–180 μM). The Michaelis–Menten constant (Km) and maximum velocity (Vmax) were determined using the PRISM v.10.0.0 software. To evaluate HCHO reduction in the presence of THF, reactions were conducted with 5 mM TMAO and varying concentrations of THF ranging from 0 to 150 µM.

Revised version: Enzyme activity was assessed by measuring HCHO production from TMAO breakdown as described previously with minor modifications.[23] Briefly, enzymatic activity was evaluated in multiple buffering systems spanning a range of pH values. Maximal activity was observed in 10 mM MES buffer (pH 6.0) supplemented with 150 mM NaCl; this buffer was subsequently used for all kinetic analyses. Standard activity assays were performed in triplicate at room temperature. Each 50 μL reaction contained 2.5 μg of purified TDM in MES buffer and was initiated by the addition of TMAO to a final concentration of 10 mM. Reactions were incubated for 10 min, a time point confirmed to be within the linear range based on preliminary time-course experiments. For HCHO detection, 10 μL of the reaction mixture was combined with 25 μL of freshly prepared 0.2% (w/v) Purpald reagent dissolved in 1 M NaOH and 215 μL of Milli-Q water in a 96-well plate. After incubation for 20 min at room temperature, absorbance at 540 nm was measured using a SpectraMax G5 microplate reader. Formaldehyde concentrations were calculated from a standard curve generated using analytical-grade formaldehyde (0–180 μM). The Michaelis–Menten constant (Km) and maximum velocity (Vmax) were determined using the PRISM v.10.0.0 software. To evaluate HCHO reduction in the presence of THF, reactions were conducted with 5 mM TMAO and varying concentrations of THF ranging from 0 to 150 µM.

For the analysis of DMA, CH2-THF products, and other compounds, high-performance liquid chromatography (HPLC) coupled with Mass Spectrometry (MS) was employed. The analysis was performed using an Agilent 1290 Infinity II LC system coupled with an Agilent 6470 series QQQ mass spectrometer. Waters Acquity UPLC with a bridged ethylene hybrid hydrophobic interaction chromatography (2.1 mm × 150 mm, 1.7 µm) column was used for separation. The mobile phases were as follows: (A) 90% acetonitrile, 5% isopropanol, and 5% 200 mM ammonium formate (pH 3) and (B) 90% water, 5% acetonitrile, and 5% 200 mM ammonium formate (pH 3). The LC gradient was: 0 min, 0% B; 2 min, 0% B; 7 min, 98% B; 7.5 minutes, 98% B; 8 min, 0% B; 15 min, 0% B, with a flow rate of 0.3 mL/min.

Mass spectrometric analysis was performed using multiple reaction monitoring (MRM) in the positive electrospray ionization (ESI) mode. The ESI interface settings were as follows: gas temperature, 325°C; gas flow rate, 7 L/min; nebulizer pressure, 45 psi; sheath gas temperature, 250°C; sheath gas flow rate, 7 L/min; capillary voltage, 3800 V; nozzle voltage, 1000 V; and ΔEMV voltage, 300 V. The concentrations of TMAO, DMA, THF, and CH2-THF were determined using internal deuterated standards (TMAO-d9 and DMA-d6), and the data were analyzed using the Agilent MassHunter Quantitative Analysis (Version B.08.00).

Cryo-EM data processing and structure determination

Earlier version: Image processing and structure determination involve several steps. Motion correction and image summation were conducted using MotionCor,[25] and defocus values were estimated using Gctf.[26] Particle picking was performed using the DoG Picker,[27] followed by initial reference-free 2D classification in CryoSPARC v.4.2.0.[28] Representative 2D class averages were selected, and particle refinement was performed through multiple rounds of 2D classification. Subsequently, an initial ab initio reconstruction was performed. The particles underwent 3D classification into five classes using a low-pass filtered initial reconstruction at 20 Å as a reference model. Multiple rounds of 3D classification were performed, and the optimal subset displaying clear structural features underwent heterogeneous refinement over six rounds, resulting in a high-quality, refined subset. These refined particles underwent homologous refinement, yielding a map with global resolution indicated by a Fourier shell correlation (FSC) of 0.143. The final volume map was computed with C2 symmetry using CryoSPARC, and local resolution was assessed using Phenix. The initial apo model of the TDM was built in Coot v0.8.9,[29] referencing the alphafold structure.[30] Monomeric TDM structures were fitted into cryo-EM electron density maps using UCSF ChimeraX,[31] with subsequent iterative manual adjustments and rebuilding in Coot, including modeling of the N-terminal domain. Manual refinement was based on the electron density map quality in Coot, followed by real-space refinement in Phenix.[32] Validation included the analysis of FSC curves to assess model-map agreement and the evaluation of atomic model geometry using MolProbity.[33] Detailed refinement statistics are provided in Table 1. The pore-lining surfaces and receptor channels were analyzed using the MOLE software.[34] Visual representations were created using UCSF ChimeraX and PyMOL.[35]

Revised version: Image processing and structure determination involve several steps. Motion correction and image summation were conducted using MotionCor,[24] and defocus values were estimated using Gctf.[25] Particle picking was performed using the DoG Picker,[26] followed by initial reference-free 2D classification in CryoSPARC v.4.2.0.[27] Representative 2D class averages were selected, and particle refinement was performed through multiple rounds of 2D classification. Subsequently, an initial ab initio reconstruction was performed. The particles underwent 3D classification into five classes using a low-pass filtered initial reconstruction at 20 Å as a reference model. Multiple rounds of 3D classification were performed, and the optimal subset displaying clear structural features underwent heterogeneous refinement over six rounds, resulting in a high-quality, refined subset. These refined particles underwent homologous refinement, yielding a map with global resolution indicated by a Fourier shell correlation (FSC) of 0.143. The final volume map was computed with C2 symmetry using CryoSPARC, and local resolution was assessed using Phenix. The initial apo model of the TDM was built in Coot v0.8.9,[28] referencing the alphafold structure.[29] Monomeric TDM structures were fitted into cryo-EM electron density maps using UCSF ChimeraX,[30] with subsequent iterative manual adjustments and rebuilding in Coot, including modeling of the N-terminal domain. Manual refinement was based on the electron density map quality in Coot, followed by real-space refinement in Phenix.[31] Validation included the analysis of FSC curves to assess model-map agreement and the evaluation of atomic model geometry using MolProbity.[32] Detailed refinement statistics are provided in Table 1. The pore-lining surfaces and receptor channels were analyzed using the MOLE software.[33] Visual representations were created using UCSF ChimeraX and PyMOL.[34]

Molecular Dynamics simulations of enzyme tunnelling

Earlier version: Coarse-grained (CG) MD simulations of TDM were conducted using the GROMACS 2021.5 and Martini3 force fields with the Gō model.[36] The CG coordinates were generated from the TDM:DMA: HCHO structure via martinize2,[37] with the secondary structure assigned using DSSP.[38] The Gō contact map was created using the GōContactMap server with a Cα–Cα distance cutoff of 0.3– 1.1 nm and Gō potential of 9.414 kJ·mol−1.[36]

Revised version: Coarse-grained (CG) MD simulations of TDM were conducted using the GROMACS 2021.5 and Martini3 force fields with the Gō model.[35] The CG coordinates were generated from the TDM:DMA: HCHO structure via martinize2,[36] with the secondary structure assigned using DSSP.[37] The Gō contact map was created using the GōContactMap server with a Cα–Cα distance cutoff of 0.3–1.1 nm and Gō potential of 9.414 kJ·mol−1.[35]

Earlier version: The simulations employed periodic boundary conditions, with short-range interaction cutoff at 1.1 nm and long-range electrostatics handled by the reaction-field method.[39] The system was maintained at 316.16 K and a pressure of 1 bar. Initial energy minimization was performed until the forces were below 10 kJ·mol−1·nm−1, followed by solvation using INSANE[40] and addition of 0.15 M NaCl. The system contained 1324 Na+/Cl− beads and 110,246 water beads in the simulation box.

Revised version: The simulations employed periodic boundary conditions, with short-range interaction cutoff at 1.1 nm and long-range electrostatics handled by the reaction-field method.[38] The system was maintained at 316.16 K and a pressure of 1 bar. Initial energy minimization was performed until the forces were below 10 kJ·mol−1·nm−1, followed by solvation using INSANE[39] and addition of 0.15 M NaCl. The system contained 1324 Na+/Cl− beads and 110,246 water beads in the simulation box.

Earlier version: The system was equilibrated for 5 ns with position restraints on the backbone beads, followed by 10 ns of unrestrained equilibration. The 10 µs production run used the velocity-rescale thermostat,[41,42] and Parrinello-Rahman barostat, with a 20 fs time step and coordinates saved every 500 ps for the production run. Tunnel calculations were performed using CAVER3.0.[43] The starting points for the tunnel calculations were defined by residues near HCHO (A282, T369, and S378). A 2.5 Å probe radius and time sparsity of 1 were applied to avoid redundant tunnel generation. The detected tunnels were clustered using a threshold of 5.0 and a frame reweighting coefficient of 1. Protein stability was evaluated by calculating RMSD, RMSF, and Rg using gmx rms, gmx rmsf, and gmx gyrates, respectively.

Revised version: The system was equilibrated for 5 ns with position restraints on the backbone beads, followed by 10 ns of unrestrained equilibration. The 10 µs production run used the velocity-rescale thermostat, and Parrinello-Rahman barostat, with a 20 fs time step and coordinates saved every 500 ps for the production run. [40,41] Tunnel calculations were performed using CAVER3.0.[42] The starting points for the tunnel calculations were defined by residues near HCHO (A282, T369, and S378). A 2.5 Å probe radius and time sparsity of 1 were applied to avoid redundant tunnel generation. The detected tunnels were clustered using a threshold of 5.0 and a frame reweighting coefficient of 1. Protein stability was evaluated by calculating RMSD, RMSF, and Rg using gmx rms, gmx rmsf, and gmx gyrates, respectively.

Isothermal Calorimetry (ITC)

ITC measurements were performed using a NANO ITC system (TA Instruments) to evaluate the binding affinities of TDM and THF at 20 °C. The compound solution was equilibrated with protein buffer D, and all protein and compound solutions were prepared in the same degassed buffer before use. ITC cells were loaded with 250 μl of TDM protein (5 μM) and titrated with THF (80 μM). Titration involved injecting 1 μl of the substrate into the protein solution, followed by 19 injections of 2.45 μl each. The stirring rate was set at 200 rpm. For control experiments, the buffer was injected into the protein alone. Equilibrium association constants were calculated by fitting reference-corrected data using the binding model supplied by the manufacturer.

Energy-dispersive X-ray spectroscopy analysis (EDAX)

Ten microliters of TDM protein at a concentration of 100 µM was dropped on the silicon wafer and vacuum dried. After drying, the samples were gold-coated (∼ 10 nm) and analyzed using a CARL ZEISS EVO 50. Buffer was used as a control experiment.

Inductively coupled plasma-mass spectrometry (ICP-MS)

Earlier version: ICP-MS experiments were conducted using trace metal-grade nitric acid (HNO3, 3% v/v), which was purified in-house by sub-boiling point distillation, as the sample matrix. ICP-MS analyses were performed using an Agilent Technologies 7500 ICP-MS instrument, while ICP-OES was conducted using a Perkin Elmer Optima 5300DV Optical Emission Spectrometer. Calibration standards were freshly prepared by diluting the stock solutions (Sigma-Aldrich) in 3% HNO3 and doubly deionized water. Approximately 2.4 mg of protein was diluted in a 3% HNO3 matrix for metal analysis, and zinc concentrations were measured using the emission line at 213.857 nm.

Revised version: ICP-MS experiments were conducted using trace metal-grade nitric acid (HNO3, 3% v/v), which was purified in-house by sub-boiling point distillation, as the sample matrix. ICP-MS analyses were performed using an Agilent Technologies 7500 ICP-MS instrument, while ICP-OES was conducted using a Perkin Elmer Optima 5300DV Optical Emission Spectrometer. Protein samples at approximately 1.0 mg/mL, together with buffer-only controls, were initially screened by ICP–MS using a stitched wide-window survey scan to approximate a full metal mass scan for the detection of metal-associated isotopes, while also identifying predominant elemental species and potential spectral interferences. Calibration standards were freshly prepared by diluting the stock solutions (Sigma-Aldrich) in 3% HNO3 and doubly deionized water. Approximately 2.4 mg of protein was diluted in a 3% HNO3 matrix for metal analysis, and zinc concentrations were measured using the emission line at 213.857 nm.